Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Simon Heming"
COLLABORATORS = "Stephen Schaumann"
```

---

# Introduction

In this exercise you will be learning the basics of programing with python and Jupyter (formerly iPython) Notebooks (Tasks 1 and 2) and algorithmic thinking (Tasks 3 and 4).

The exercises of this lecture will be using Jupyter in combination with the nbgrader extension. Each exercise will contain tasks that are to be answered in free text or with source code. The source code should adhere to PEP-8 style guidelines and should be commented so that the tutor understands what you are doing (points may be deducted in both cases). In some cases, a source answer cell might be followed by a Python unittest. If this test fails, you might not get points. On the other hand, if it succeeds you might not either (in case you try to cheat the test):

```python
	def pythagoras(a, b):
		return 5

	assert pythagoras(3, 4) == 5
```

# Introduction to Jupyter Notebooks and Python (7 Pts.)
 
> The Jupyter Notebook is an open-source web application that allows you to create and share documents that contain live code, equations, visualizations and narrative text. Uses include: data cleaning and transformation, numerical simulation, statistical modeling, data visualization, machine learning, and much more. - https://jupyter.org/index.html

1. Install Python 3.5 or higher with Jupyter (<http://jupyter.org/install>)
	- Suggestion for Windows: Use Anaconda
	- Suggestion for Linux-OSes: Use a virtual environment and pip
1. (Optional but highly recommended) Work through chapters 3, 4, 5, 7.2, 8 und 9.3 to get a first look at python. Chapters 4.7, 5.1.3, 5.1.4, 5.6, 5.7 and 5.8 may be skipped. Furthermore, have a look at <docs.python.org/library/index.html> to get an overview of built-in packages. Use this notebook as a scratch pad by inserting a new cell below.



## Answer the following questions
(Answer in the markdown cell, refer to code in the following code cell.)
1. In which module can the function `sqrt()` be found?
1. What happens when `sqrt()` is called with a negative number?
1. Implement two functions `mysqrt()` that prints *`mysqrt()` does not work with negative numbers, you fool.* when it is called with a negative number, and the result in the other case.
    * Use `if: … else:`
    * Use `try: … except:`
1. Implement a loop that prints `i mod 5 = <result>` for `i = [-10..10]` and explain the result of the modulo operation.
1. When should a string be enclosed in triple quotes?
1. What is the difference between the classes `dict` and `list`?
1. What is the purpose of the `__init__()` function of a class and how is it used?


## Answer the following questions
(Answer in the markdown cell, refer to code in the following code cell.)
1. In which module can the function `sqrt()` be found? -> math module `from math import sqrt`
1. What happens when `sqrt()` is called with a negative number? -> Value error
1. Implement two functions `mysqrt()` that prints *`mysqrt()` does not work with negative numbers, you fool.* when it is called with a negative number, and the result in the other case. 
    * Use `if: … else:`
    * Use `try: … except:`
1. Implement a loop that prints `i mod 5 = <result>` for `i = [-10..10]` and explain the result of the modulo operation.
1. When should a string be enclosed in triple quotes? -> when the string contains single quotes e.g.: `'''Cannot use 'x' as parameter'''`
1. What is the difference between the classes `dict` and `list`?
1. What is the purpose of the `__init__()` function of a class and how is it used? -> init the first call when the class is used. Should be used for all member variables. Can change the constructor of the class e.g.: Point2D in code


<!-- grade_id: cell-1f1f5d99b4ea02f9, points: 7, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-aa57d5b55494c658, points: 7, grade: true, solution: true -->
```python
# YOUR CODE HERE
from math import sqrt

def mysqrt(number):
    if number < 0:
        return 'mysqrt() does not work with negative numbers, you fool.'
    else:
        try:
            return sqrt(number)
        except Exception as e:
            print('error: {}'.format(e))

def mod_explained():
    for i in range(-10, 10, 1):
        print(i % 5)
mod_explained()

class Point2D():
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def __str__(self):
        return 'Point2D: x = {}, y = {}'.format(self.x, self.y)
p = Point2D(5, 4)
print(p)
```

Output:
```
0
1
2
3
4
0
1
2
3
4
0
1
2
3
4
0
1
2
3
4
Point2D: x = 5, y = 4
```

# Sieve of Eratosthenes (6 Pts.)
Implement the [Sieve](https://en.wiktionary.org/wiki/sieve) of Eratosthenes used to find prime numbers (To be found on the internet). Make the tests pass.

```python
sieve(n)  # print all primes below n.
```

<!-- grade_id: sieve, points: 0, grade: false, solution: true -->
```python
from math import sqrt
from math import ceil


def sieve(n):
    """Compute and return all prime numbers below n+1"""
    # YOUR CODE HERE
    if n < 2:
        raise ValueError
    primes = [x for x in range(2, n + 1)]
    for x in primes:
        for i, prime in enumerate(primes):
            if prime % x is 0 and prime is not x:
                del primes[i]
    return primes
print(len(sieve(500)))
```

Output:
```
95
```

<!-- grade_id: test_sieve_values, points: 3, grade: true, solution: false -->
```python
assert sieve(2) == [2]
assert sieve(3) == [2, 3]
assert sieve(4) == [2, 3]
assert sieve(5) == [2, 3, 5]

```

<!-- grade_id: test_sieve_invalid_input, points: 3, grade: true, solution: false -->
```python
try:
    sieve(0)
except ValueError:
    pass
else:
    raise AssertionError("Did not raise")
    
try:
    sieve(-4)
except ValueError:
    pass
else:
    raise AssertionError("Did not raise")

```

# Friday 13th (8 Pts.)

Once every few month there will be a Friday 13th, but how often does this event actually occur?

1. How many possible weekday - date combinations are possible? (Don’t forget leap years)
1. Determine how many Friday 13th you have experienced in your life


<!-- grade_id: Antwort1, points: 8, grade: true, solution: true -->
YOUR ANSWER HERE

# Stable Glasses (12 Pts.)
Glases in pubs are prone to breaking. Therefore, more stable glasses are being developed. To evaluate the stability of a new kind of glas, the following experiment is being done: Platforms that allow a controlled fall of a glass can be attached to a 3 m high stand with possible platform positions every 10 cm. The goal of this experiment is to find the highest possible fall height for a glass not to shatter with as few tries as possible. One try is defined as one fall of a glass off a platform. Another constraint is the number of available glasses which means only a certain amount of glasses may break.

1. How do you determine the height of fall in case only one glass is at hand? How many tries are needed at most (explain). From an algorithmic standpoint, this is the worst case, for pubs the best case.
1. It is possible to use fewer tries to reach the goal if more glasses may break. What is the process in this case, how many tries are needed at most and how many glasses are broken at most?
1. A compromise is to allow two glasses to break. Which process minimizes the number of tries in the worst case scenario in this test case and how many tries will be needed in the worst case? (Hint: less than 15)
1. Generalize your result from 3. s.t. the stand has `n` possible platform position. The process remains the same but the number of tries in the worst case is now a function `worst_case(n)`. Determine the function! Since the the number of tries is aways an integer, the function is a step function: The results for certain intervals may be the same until it increases by one only to be constant again for a few more steps. For `n=30`, the result should be the same as in 3.

<!-- grade_id: Antwort2, points: 10, grade: true, solution: true -->
3 m high with platforms every 10 cm -> total of 30 possible platforms
1. to determine the maximum height the platform is moved down 10 cm per experiment. this would mean a total of 30 tries at maximum, because the glass will never break
1. Divide and conquer approach. At the beginning the minimum height is 0 m maximum height is 3 m. (ceiling while testing) Test if the glass breaks at minheight + maxheight / 2. If it does maxheight = maxheight / 2. If it doesnt minheight = maxheight / 2. repeat. The maximum number of tries are 5. The maximum number of glasses that would be broken this way are 5 (e. g. if the glass already breaks at 10 cm)
1. 

<!-- grade_id: cell-a2426d9a7636bfde, points: 0, grade: false, solution: true -->
```python
from math import ceil, sqrt, floor

def worst_case(n):
    """Compute worst case."""
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-f80d4f8640c3d48e, points: 2, grade: true, solution: false -->
```python
try:
    worst_case(0)
except ValueError:
    pass
else:
    raise AssertionError("Did not raise")

```

```python

```