Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Simon Heming"
COLLABORATORS = "Stephen Schaumann"
```

---

#  Einleitung
In dieser Aufgabe soll Ihr Verständnis von Sortieralgorithmen gestärkt werden. Stellen Sie sicher, dass sie das Archiv samt Test-Dateien heruntergeladen und in einen Unterordner `test_cases` extrahiert haben (Quelle: http://blog.quibb.org/2008/11/sort-optimization/). Öffnen Sie den Ordner und schauen Sie sich den Schema an, nachdem die Dateien benannt sind.




# Sortierverfahren implementieren (6⋅4 Pkt.)
Implementieren Sie folgende Sortieralgorithmen.

- Selectionsort
- Insertionsort
- Quicksort
- Bubblesort
- Mergesort
- Heapsort

Die Sortieralgorithmen sollen eine sortierte Kopie der ursprünglichen Liste zurück geben und das Original nicht verändern (damit das Verhalten kongruent zu `sorted()`) ist.


<!-- grade_id: cell-6d812715dbe509ce, points: 0, grade: false, solution: true -->
```python
import numpy as np
def selectionsort(data):
    s = []
    u = list(data)
    
    while(len(u) > 0):
        minimum = u[0]
        index = 0
        for i in range(1, len(u)):
            if u[i] < minimum:
                minimum = u[i]
                index = i
        u[index] = u[0]
        u[0] = minimum
        s.append(u[0])
        u = u[1:]
    return s

def insertionsort(data):
    ret_data = list(data)
    for i in range(1, len(ret_data)):
        compare = ret_data[i]
        j = i - 1
        while(j >= 0 and ret_data[j] > compare):
            ret_data[j + 1] = ret_data[j]
            j -= 1
        ret_data[j + 1] = compare
    return ret_data

def bubblesort(data):
    ret_data = list(data)
    for i in range(len(ret_data)):
        for j in range(len(ret_data)-1):
            if ret_data[j] > ret_data[j + 1]:
                switch = ret_data[j]
                ret_data[j] = ret_data[j + 1]
                ret_data[j + 1] = switch
    return ret_data

def quicksort(data):
    ret_data = list(data)
    if len(ret_data) == 1:
        return ret_data[0]
    median = ret_data[len(ret_data)//2] -1 if ret_data[len(ret_data)//2] -1 > 0 else 0
    median = np.mean(ret_data)
    smaller = []
    bigger = []
    for i in range(len(ret_data)):
        if ret_data[i] > median:
            bigger.append(ret_data[i])
        else:
            smaller.append(ret_data[i])
    print('median: {}, smaller: {}, bigger: {}'.format(median, smaller, bigger))
    quicksort(smaller)
    quicksort(bigger)
    for i in range(len(smaller) - 1):
        if smaller[i] > smaller[i + 1]:
            tmp = smaller[i]
            smaller[i] = smaller[i + 1]
            smaller[i + 1] = tmp
    for i in range(len(bigger) - 1):
        if bigger[i] > bigger[i + 1]:
            tmp = bigger[i]
            bigger[i] = bigger[i + 1]
            bigger[i + 1] = bigger
    return ret_data

def heapsort(data):
    # YOUR CODE HERE
    raise NotImplementedError()

def mergesort(data):
    sorted_list = []
    if len(data) == 1:
        return data
    index = len(data)//2
    data1 = data[:index]
    data2 = data[index:]
    print(data1)
    print(data2)
    mergesort(data1)
    mergesort(data2)
    for i in range(len(data1)):
        for j in range(len(data2)):
            if data1[i] > data2[j]:
                pass
    

#mergesort(da)

```

```python
data = [1,23,452,74,21,6,9,324,675,34,21,0]
quicksort(data)
```

Output:
```
median: 136.66666666666666, smaller: [1, 23, 74, 21, 6, 9, 34, 21, 0], bigger: [452, 324, 675]
median: 21.0, smaller: [1, 21, 6, 9, 21, 0], bigger: [23, 74, 34]
median: 9.666666666666666, smaller: [1, 6, 9, 0], bigger: [21, 21]
median: 4.0, smaller: [1, 0], bigger: [6, 9]
median: 0.5, smaller: [0], bigger: [1]
median: 7.5, smaller: [6], bigger: [9]
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], bigger: []
median: 21.0, smaller: [21, 21], 
```

<!-- grade_id: cell-848f5caf61f83589, points: 0, grade: true, solution: false -->
```python
import random

SAMPLE_LIST = random.sample(range(100), 30)

def test_algorithm(algorithm):
    sample_list = SAMPLE_LIST[:]
    assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
    assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)
```

<!-- grade_id: cell-87b11519e17f1c63, points: 4, grade: true, solution: false -->
```python
test_algorithm(selectionsort)
```

<!-- grade_id: cell-24cfee0a571db291, points: 4, grade: true, solution: false -->
```python
test_algorithm(insertionsort)
```

<!-- grade_id: cell-78ffe6f52e47fc47, points: 4, grade: true, solution: false -->
```python
test_algorithm(bubblesort)
```

<!-- grade_id: cell-e60d068bfdb9cdde, points: 4, grade: true, solution: false -->
```python
test_algorithm(quicksort)
```

Output:
```
... last 1 frames repeated, from the frame below ...

<ipython-input-109-17322f50927f> in quicksort(data)
     83         else:
     84             smaller.append(data[i])
---> 85     quicksort(smaller)
     86     quicksort(bigger)
     87 

RecursionError: maximum recursion depth exceeded in comparison
```

<!-- grade_id: cell-037297fb44664ec1, points: 4, grade: true, solution: false -->
```python
test_algorithm(heapsort)
```

<!-- grade_id: cell-96e6149464b940df, points: 4, grade: true, solution: false -->
```python
test_algorithm(mergesort)
```

# Einlesen von Dateien (4 Pkt.)
Schreiben sie eine Funktion `load_test_cases(path)`, die zu den zur Verfügung gestellten Testfällen jeweils die Datei mit 10e3 Einträgen lädt.

<!-- grade_id: cell-d64d799097780639, points: 0, grade: false, solution: true -->
```python
from glob import glob


TEST_CASE_PATH = "/home/heming/Downloads/test_files/"

def load_test_cases(path):
    data_dir = glob(path + '/*')
    find = '.10000.txt'
    data = {}
    for file in data_dir:
        if find in file:
            file_name = file.split('/')[-1]
            key_name = '.'.join(file_name.split('.')[0:-2])
            txt_data = []
            with open(file) as f:
                for line in f:
                    txt_data.append(int(line))
            data[key_name] = txt_data
    return data
```

```python

```

<!-- grade_id: cell-6b64b68014341d9a, points: 4, grade: true, solution: false -->
```python
import unittest

unittest_test_case = unittest.TestCase('__init__')

loaded_test_cases = load_test_cases(TEST_CASE_PATH)
available_test_cases = ['1245.repeat', 'end.0', 'identical', 'med.3.killer', 'rand.dups', 'rand.no.dups', 'rand.steps', 
                        'rev.ends', 'rev.partial', 'rev.saw', 'reverse', 'seq.0.is.1000', 'seq.partial', 'seq.saw', 'sequential']

unittest_test_case.assertCountEqual(loaded_test_cases.keys(), available_test_cases)

for key, value in loaded_test_cases.items():
    assert len(value) == 10e3, "Number of testcases mismatch for {}".format(key)
```

# Sortierverfahren testen (12 Pkt.)
 Vergleichen Sie die Sortierverfahren untereinander und mit der Python Funktion `sorted()`
 - Benutzen sie `result = %timeit -n 100 -r 3 -o` um die Ausführung zu messen und das Ergebnis in einer Variablen speichern zu können. Auf das beste Ergebnis können Sie mit `result.best` zugreifen.
 - Gruppieren Sie die Ergebnisse nach Testfall.
 - Testen sie folgende „Splices“ Ihres Testdatensatzes: `x_values = [10, 20, 50, 100, 200, 500, 1000, 2000, 5000, 10000]`
 - Stellen sie die Ergebnisse mittels Tabelle und Graphen mit Titel und Achsenbeschriftung dar (Matplotlib).
 - Bewerten Sie das Ergebnis

One Image with multiple algorithms per File

<!-- grade_id: cell-584a475079bfa4bd, points: 12, grade: true, solution: true -->
```python
# for quicksort
import sys
sys.setrecursionlimit(20000)

for key, value in loaded_test_cases.items():
    werte = value[:1000]
    #print(werte)
    result = %timeit -n 100 -r 3 -o bubblesort(werte)
    print(result.best)
        
```

Output:
```
96.1 ms ± 328 µs per loop (mean ± std. dev. of 3 runs, 100 loops each)
0.09575852424999538
123 ms ± 97 µs per loop (mean ± std. dev. of 3 runs, 100 loops each)
0.12320512444999622
<magic-timeit> in inner(_it, _timer)

<ipython-input-136-c31fb0f9c4b7> in bubblesort(data)
     66     for i in range(len(ret_data)):
     67         for j in range(len(ret_data)-1):
---> 68             if ret_data[j] > ret_data[j + 1]:
     69                 switch = ret_data[j]
     70                 ret_data[j] = ret_data[j + 1]

KeyboardInterrupt:
```

```python

```