Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Stephen Schaumann"
COLLABORATORS = "Simon Heming"
```

---

# Aufgabe 1 - Treaps (40 Punkte)

Bei selbstbalancierenden Bäumen versucht man, die Höhe des Baumes zu minimieren. Dies ist aber gar nicht das eigentliche Ziel der Optimierung: In Wirklichkeit will man die Zugriffszeit auf die Elemente minimieren. Wenn die Schlüssel mit sehr unterschiedlicher Häufigkeit abgefragt werden, ist ein balancierter Baum dafür nicht die optimale Lösung. Stattdessen will man häufige Schlüssel nahe der Wurzel des Baumes abspeichern, auch wenn seltene Schlüssel dadurch in tieferen Ebenen landen als beim balancierten Baum. Dies wird durch den sogenannten Treap erreicht, der die Eigenschaften von Suchbäumen und Heaps verbindet. Neben dem Schlüssel hat hier jedes Element auch eine Priorität, und Elemente mit hoher Priorität liegen nahe der Wurzel des Baumes. Dies erreicht man, indem jeder Teilbaum zwei Eigenschaften erfüllt:  

>(1) Sortierung nach Schlüsseln: Alle Elemente im linken Teilbaum haben kleinere Schlüssel als die Wurzel des Teilbaums, alle Elemente im rechten Teilbaum größere (Suchbaumbedingung).  

>(2) Sortierung nach Prioritäten: Kein Element im linken oder rechten Teilbaum hat höhere Priorität als die Wurzel des Teilbaums (Heap-Bedingung).  

Die Erfinder der Treap-Datenstruktur haben gezeigt, dass beide Eigenschaften gleichzeitig erfüllbar sind. Der grundlegende Einfügealgorithmus besteht aus den folgenden Schritten:  

>1. Füge ein neues Element entsprechend seines Schlüssels ein wie in einen unbalancierten Baum (vergleiche **insert()** aus Übung 5). Damit ist Bedingung (1) erfüllt.  

>2. Wenn die Priorität des neuen Elements höher ist als die seines Vaterknotens: Führe eine Rota-tion aus, die das neue Element eine Ebene nach oben bewegt. Wenn das neue Element das linke Kind ist, muss eine Rechtsrotation ausgeführt werden und umgekehrt.  

>3. Wiederhole Schritt 2 auf der nächsthöheren Ebene bis entweder Bedingung (2) erfüllt oder das neue Element zur Wurzel des Treaps geworden ist. Da Bedingung (1) invariant gegenüber Rotationen ist, entsteht dadurch immer ein gültiger Treap.  

Für die Festlegung der Prioritäten gibt es mehrere Möglichkeiten:  

>(A)	Wenn man konstante Prioritäten verwendet, entsteht ein unbalancierter Baum.  

>(B)	Wenn man Zufallszahlen verwendet, entsteht ein näherungsweise balancierter Baum.  

>(C)	Wenn man die Wahrscheinlichkeiten verwendet, mit der auf die Schlüssel zugegriffen wird, entsteht ein näherungsweise zugriffsoptimaler Baum.  

>(D)	Wenn man bei jedem Zugriff auf ein Element dessen Priorität inkrementiert (und den Baum umstrukturiert, falls Bedingung (2) nicht mehr erfüllt ist), passt sich der Baum automatisch an variierende Zugriffsmuster an, indem häufig benutzte Schlüssel nach oben wandern. Beim ersten Einfügen wird die Priorität mit 1 initialisiert.  

Lösen Sie folgende Aufgaben:  

a) **(4 Points)**  
Fügen Sie die folgenden Zeichen mit ihren Prioritäten in einen leeren Treap ein und beantworton Sie die Fragen:

>(A, 17), (F, 22), (Q, 15), (K, 25), (R, 14), (B, 35), (X, 27), (M, 24), (L, 42), (T, 30)

1. Geben Sie die Länge des Treap an (0 heißt, dass sich nur ein Knoten im Treip befindet)
2. Welches Zeichen ist die Wurzel des Treap
3. Wie viele Blätter hat der Treap
4. Welches Zeichen (falls vorhanden) ist der rechte Kind-Knoten von M

<!-- grade_id: cell-2, points: 4, grade: true, solution: true -->
1. Die Länge des Treap ist 4
2. Die Wurzel ist L (es hat die höchste Priorität mit 42)
3. Der Treap hat 4 Blätter
4. Der rechte Kind-Knoten von M ist Q

b) **(2 Points)**  
Löschen Sie (L, 42) aus dem Treap Und beantworten Sie dieselben Fragen.

<!-- grade_id: cell-4, points: 2, grade: true, solution: true -->
1. Die Länge des Treap ist 5
2. Die Wurzel ist B (es hat die höchste verbleibende Priorität, 35)
3. Der Treap hat 4 Blätter
4. Der rechte Knoten von M ist Q

c)	**(2 Punkte)**   
Was ist der Unterschied zwischen den Eigenschaften des Binären-Suchbaums und der den Min-Heap? Kann die Min-Heap Eigenschaft ausgenutzt werden um die Schlüssel eines Baums mit $n$ Knoten in $O\left(n\right)$ auszugeben? Begründen Sie ihre Antwort.

<!-- grade_id: cell-6, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE


d)	**(6 Punkte)**   
Implementieren Sie die aus der Vorlesung bekannten Operationen  
  >*newroot = rotateLeft(rootnode)*  
  >*newroot = rotateRight(rootnode)*
  
**Hinweis**   
Durch eine Rotation ändern wir die Struktur der Zeiger. Dies ist eine lokale Operation im Suchbaum durch die die Binäre-Suchbaum-Eigencshaft erhalten bleibt. Wenn wir beispielsweise eine Linksrotation an Knoten $x$ durchführen nehmen wir an, dass das rechte Kindelement von $x$, $y$ nicht leer ist; $x$ kann jeder Knoten im Baum sein, dessen rechtes Kind-Element nicht leer ist. Die Linksrotation dreht sich um die Verbindung von $x$ zu $y$. Dadurch wird $y$ die neue Wurzel des Teilbaums wobei $x$ das linke Kindelement von $y$ wird und das linke Kindelement von $y$ das rechte Kindelement von $x$ wird.

```python
class Node:
    def __init__(self, key, priority):
        self.key = key
        self.priority = priority
        self.left = self.right = None
    
    # for test
    def left_key(self):
        return 0 if self.left is None else self.left.key
    def right_key(self):
        return 0 if self.right is None else self.right.key
```

<!-- grade_id: cell-9, points: 0, grade: false, solution: true -->
```python
def rotateLeft(rootnode):
    if rootnode.right == None:
        raise ValueError()
    elif rootnode.left == None:
        subtree = rootnode.right
        rootnode.right = None
        subtree.left = rootnode
    else:
        subtree = rootnode.right
        rootnode.right = subtree.left
        subtree.left = rootnode
    

def rotateRight(rootnode):
    if rootnode.left == None:
        raise ValueError()
    elif rootnode.right == None:
        subtree = rootnode.left
        rootnode.left = None
        subtree.right = rootnode
    else:
        subtree = rootnode.left
        rootnode.left = subtree.right
        subtree.right = rootnode
```

<!-- grade_id: cell-10, points: 2, grade: true, solution: false -->
```python
# --- test left rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n1.right = n2
assert n1.left_key() == 0
assert n1.right_key() == 2
assert n2.left_key() == 0
assert n2.right_key() == 0
rotateLeft(n1)
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 0
# --- test right rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n2.left = n1
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 0
rotateRight(n2)
assert n1.left_key() == 0
assert n1.right_key() == 2
assert n2.left_key() == 0
assert n2.right_key() == 0
```

<!-- grade_id: cell-11, points: 4, grade: true, solution: false -->
```python
# --- test left rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n3 = Node(3,3)
n4 = Node(4,4)
n5 = Node(5,5)
n2.left = n1
n2.right = n4
n4.left = n3
n4.right = n5
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 4
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 3
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
rotateLeft(n2)
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 3
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 2
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
# --- test right rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n3 = Node(3,3)
n4 = Node(4,4)
n5 = Node(5,5)
n2.left = n1
n2.right = n3
n4.left = n2
n4.right = n5
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 3
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 2
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
rotateRight(n4)
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 4
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 3
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
```

e)	**(24 Punkte)**   
Implementieren Sie die Klassen **RandomTreap** und **DynamicTreap**, die die Prioritäten nach Variante (B) bzw. (D) festlegen. Gehen Sie dabei von der Klasse **SearchTree** aus Übung 5 aus. Die Node-Klasse muss um ein Attribut priority erweitert werden, das value-Attribut kann dafür wegfallen, weil es in dieser Übung nicht benötigt wird. Die Funktion **insert()** realisiert die Treap-Semantik: Ist der key noch nicht im Treap enthalten, wird er nach der Suchbaumregel eingefügt und der Baum eventuell entsprechend der Prioritäten umstrukturiert. Ist der key bereits vorhanden, passiert beim RandomTreap nichts, während beim DynamicTreap die Prio-rität um Eins erhöht und der Baum gegebenenfalls umstrukturiert wird.   
  
**Bemerkung**  
Die Funktion `insert()` von `RandomTreap`, hat ein optionales Argument `priority`. Wird dieses Argument weggelassen wird die Priorität zufällig durch Pythons `random.random()` Funktion festgelegt.

<!-- grade_id: cell-13, points: 0, grade: false, solution: true -->
```python
import random
### For test purpose, please use random.random() to set the priority of each node.

class RandomTreap:
    def __init__(self):
        self.root = None

    def insert(self, key, priority=None):
        if priority is None:
            priority = random.random()
        n = Node(key, priority)
        if self.root == None:
            self.root = n
        else:
            compare = self.root
            while True:
                if compare.key == n.key:
                    break
                elif compare.key > n.key:
                    if compare.left is None:
                        compare.left = n
                        break
                    else:
                        compare = compare.left
                elif compare.key < n.key:
                    if compare.right is None:
                        compare.right = n
                        break
                    else:
                        compare = compare.right
        self.fix_prio(self.root)
    
    def fix_prio(self, subroot, last=None):
        #print('subroot key: {}'.format(subroot.key))
        if subroot.left is not None:
            self.fix_prio(subroot.left, subroot)
        if subroot.right is not None:
            self.fix_prio(subroot.right, subroot)
        if subroot.left is not None:
            if subroot.left.priority > subroot.priority:
                self.rotateRight(subroot, last)
        if subroot.right is not None:
            if subroot.right.priority > subroot.priority:
                self.rotateLeft(subroot, last)
                
    def rotateLeft(self, rootnode, last=None):
        if rootnode.right == None:
            raise ValueError()
        elif rootnode.left == None:
            subtree = rootnode.right
            rootnode.right = None
            subtree.left = rootnode
        else:
            subtree = rootnode.right
            rootnode.right = subtree.left
            subtree.left = rootnode
        if last == None:
            self.root = subtree
        else:
            if last.left == rootnode:
                last.left = subtree
            if last.right == rootnode:
                last.right = subtree
    

    def rotateRight(self, rootnode, last=None):
        if rootnode.left == None:
            raise ValueError()
        elif rootnode.right == None:
            subtree = rootnode.left
            rootnode.left = None
            subtree.right = rootnode
        else:
            subtree = rootnode.left
            rootnode.left = subtree.right
            subtree.right = rootnode
        if last == None:
            self.root = subtree
        else:
            if last.left == rootnode:
                last.left = subtree
            if last.right == rootnode:
                last.right = subtree
            
    def remove(self, key):
        self.remove_(self.root, key)
    
    def remove_(self, node, key, last=None):
        if node.key == key:
            self.handle_deletion(node, key, last)
        elif node.key > key:
            if node.left == None:
                print('key: {} is not part of the treap'.format(key))
            else:
                self.remove_(node.left, key, node)
        elif node.key < key:
            if node.right == None:
                print('key: {} is not part of the treap'.format(key))
            else:
                self.remove_(node.right, key, node)
    
    def handle_deletion(self, node, key, last):
        if last == None:
            # root deletion
            if node.right is None and node.left is None:
                self.root = None
            elif node.right is None:
                self.root = node.left
            elif node.left is None:
                self.root = node.right
            else:
                leftmost, parent = node.right, node
                while leftmost.left is not None:
                    leftmost, parent = leftmost.left, leftmost
                if leftmost == node.right:
                    leftmost.left = node.left
                else:
                    if leftmost.right is not None:
                        parent.left = leftmost.right
                    else:
                        parent.left = None   
                    leftmost.left, leftmost.right = node.left, node.right
                self.root = leftmost
        else:
            # other deletions
            if node.right is None and node.left is None:
                if last.right == node:
                    last.right = None
                if last.left == node:
                    last.left = None
            elif node.right is None:
                if last.right == node:
                    last.right = node.left
                if last.left == node:
                    last.left = node.left
            elif node.left is None:
                if last.right == node:
                    last.right = node.right
                if last.left == node:
                    last.left = node.right
            else:
                leftmost, parent = node.right, node
                while leftmost.left is not None:
                    leftmost, parent = leftmost.left, leftmost
                if node.right is not leftmost:
                    if leftmost.right is not None:
                        parent.left = leftmost.right
                    else:
                        parent.left = None
                    leftmost.left, leftmost.right = node.left, node.right
                    if last.right == node:
                        last.right = leftmost
                    elif last.left == node:
                        last.left = leftmost
                else:
                    if last.right == node:
                        last.right = leftmost
                    elif last.left == node:
                        last.left = leftmost
                    leftmost.left = node.left
        self.fix_prio(self.root)
# Traverse
def preOrderTraverse(root):
    if root is not None:
        print ('key', root.key, 'priority', root.priority)
        preOrderTraverse(root.left)
        preOrderTraverse(root.right)
            
def inOrderTraverse(root):
    if root is not None:
        inOrderTraverse(root.left)
        print ('key', root.key, 'priority', root.priority)
        inOrderTraverse(root.right)

def postOrderTraverse(root):
    if root is not None:
        postOrderTraverse(root.left)
        postOrderTraverse(root.right)
        print ('key', root.key, 'priority', root.priority)

```

```python
# Check
def check(root):
    if root is not None:
        if root.left is not None:
            assert root.key >= root.left.key, "------------------\n!!! Error !!! Not a treap: the BST condition is violated!\n------------------"
            assert root.priority >= root.left.priority, "------------------\n!!! Error !!! Not a treap: the heap condition is violated!\n------------------"
        if root.right is not None:
            assert root.key <= root.right.key, "------------------\n!!! Error !!! Not a treap: the BST condition is violated!\n------------------"
            assert root.priority >= root.right.priority, "------------------\n!!! Error !!! Not a treap: the heap condition is violated!\n------------------"
        check(root.left)
        check(root.right)
```

<!-- grade_id: cell-14, points: 4, grade: true, solution: false -->
```python
random.seed(1)
a = RandomTreap()
a.insert(1,17) 
a.insert(6,22)
a.insert(17,15)
a.insert(11,25)
a.insert(18,14)
a.insert(2,35)
a.insert(24,27)
a.insert(13,24)
a.insert(12,42)
a.insert(20,30)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
check(a.root)
a.remove(12)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
check(a.root)
```

Output:
```
--- preOrderTraverse ---
key 12 priority 42
key 2 priority 35
key 1 priority 17
key 11 priority 25
key 6 priority 22
key 20 priority 30
key 24 priority 27
--- inOrderTraverse ---
key 1 priority 17
key 2 priority 35
key 6 priority 22
key 11 priority 25
key 12 priority 42
key 20 priority 30
key 24 priority 27
--- postOrderTraverse ---
key 1 priority 17
key 6 priority 22
key 11 priority 25
key 2 priority 35
key 24 priority 27
key 20 priority 30
key 12 priority 42
--- preOrderTraverse ---
key 2 priority 35
key 1 priority 17
key 20 priority 30
key 11 priority 25
key 6 priority 22
key 24 priority 27
--- inOrderTraverse ---
key 1 priority 17
key 2 priority 35
key 6 priority 22
key 11 priority 25
key 20 priority 30
key 24 priority 27
--- postOrderTraverse ---
key 1 priority 17
key 6 priority 22
key 11 priority 25
key 24 priority 27
key 20 priority 30
key 2 priority 35
```

<!-- grade_id: cell-15, points: 4, grade: true, solution: false -->
```python
random.seed(1)
a = RandomTreap()
a.insert(6) 
a.insert(2)
a.insert(5)
a.insert(3)
a.insert(8)
a.insert(4)
a.insert(1)
a.insert(7)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
check(a.root)
```

Output:
```
--- preOrderTraverse ---
key 2 priority 0.8474337369372327
key 1 priority 0.651592972722763
key 7 priority 0.7887233511355132
key 5 priority 0.763774618976614
key 4 priority 0.4494910647887381
key 3 priority 0.2550690257394217
key 6 priority 0.13436424411240122
key 8 priority 0.49543508709194095
--- inOrderTraverse ---
key 1 priority 0.651592972722763
key 2 priority 0.8474337369372327
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 5 priority 0.763774618976614
key 6 priority 0.13436424411240122
key 7 priority 0.7887233511355132
key 8 priority 0.49543508709194095
--- postOrderTraverse ---
key 1 priority 0.651592972722763
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 6 priority 0.13436424411240122
key 5 priority 0.763774618976614
key 8 priority 0.49543508709194095
key 7 priority 0.7887233511355132
key 2 priority 0.8474337369372327
```

<!-- grade_id: cell-16, points: 2, grade: true, solution: false -->
```python
a.insert(8)
a.insert(4)
a.insert(1)
a.insert(7)
assert a.root.key == 2
assert a.root.left.key == 1
assert a.root.right.key == 7
assert a.root.right.left.key == 5
assert a.root.right.left.left.key == 4
assert a.root.right.left.left.left.key == 3
assert a.root.right.left.right.key == 6
assert a.root.right.right.key == 8
```

<!-- grade_id: cell-17, points: 4, grade: true, solution: false -->
```python
a.remove(7)
check(a.root)
a.remove(1)
check(a.root)
a.insert(1)
check(a.root)
a.insert(7)
check(a.root)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
```

Output:
```
--- preOrderTraverse ---
key 2 priority 0.8474337369372327
key 1 priority 0.762280082457942
key 5 priority 0.763774618976614
key 4 priority 0.4494910647887381
key 3 priority 0.2550690257394217
key 8 priority 0.49543508709194095
key 7 priority 0.0021060533511106927
--- inOrderTraverse ---
key 1 priority 0.762280082457942
key 2 priority 0.8474337369372327
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 5 priority 0.763774618976614
key 7 priority 0.0021060533511106927
key 8 priority 0.49543508709194095
--- postOrderTraverse ---
key 1 priority 0.762280082457942
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 7 priority 0.0021060533511106927
key 8 priority 0.49543508709194095
key 5 priority 0.763774618976614
key 2 priority 0.8474337369372327
```

Implementieren Sie  `DynamicTreap` in der folgenden Zelle.

<!-- grade_id: cell-19, points: 0, grade: false, solution: true -->
```python
class DynamicTreap:
    def __init__(self):
        self.root = None

    def insert(self, key):
        priority = 0
        n = Node(key, priority)
        if self.root == None:
            self.root = n
        else:
            compare = self.root
            while True:
                if compare.key == n.key:
                    compare.priority += 1
                    break
                elif compare.key > n.key:
                    if compare.left is None:
                        compare.left = n
                        break
                    else:
                        compare = compare.left
                elif compare.key < n.key:
                    if compare.right is None:
                        compare.right = n
                        break
                    else:
                        compare = compare.right   
        self.fix_prio(self.root)
        
    def fix_prio(self, subroot, last=None):
        if subroot.left is not None:
            self.fix_prio(subroot.left, subroot)
        if subroot.right is not None:
            self.fix_prio(subroot.right, subroot)
        if subroot.left is not None:
            if subroot.left.priority > subroot.priority:
                self.rotateRight(subroot, last)
        if subroot.right is not None:
            if subroot.right.priority > subroot.priority:
                self.rotateLeft(subroot, last)
                
    def rotateLeft(self, rootnode, last=None):
        if rootnode.right == None:
            raise ValueError()
        elif rootnode.left == None:
            subtree = rootnode.right
            rootnode.right = None
            subtree.left = rootnode
        else:
            subtree = rootnode.right
            rootnode.right = subtree.left
            subtree.left = rootnode
        if last == None:
            self.root = subtree
        else:
            if last.left == rootnode:
                last.left = subtree
            if last.right == rootnode:
                last.right = subtree
    

    def rotateRight(self, rootnode, last=None):
        if rootnode.left == None:
            raise ValueError()
        elif rootnode.right == None:
            subtree = rootnode.left
            rootnode.left = None
            subtree.right = rootnode
        else:
            subtree = rootnode.left
            rootnode.left = subtree.right
            subtree.right = rootnode
        if last == None:
            self.root = subtree
        else:
            if last.left == rootnode:
                last.left = subtree
            if last.right == rootnode:
                last.right = subtree
            
    def remove(self, key):
        self.remove_(self.root, key)
    
    def remove_(self, node, key, last=None):
        if node.key == key:
            self.handle_deletion(node, key, last)
        elif node.key > key:
            if node.left == None:
                print('key: {} is not part of the treap'.format(key))
            else:
                self.remove_(node.left, key, node)
        elif node.key < key:
            if node.right == None:
                print('key: {} is not part of the treap'.format(key))
            else:
                self.remove_(node.right, key, node)
    
    def handle_deletion(self, node, key, last):
        if last == None:
            # root deletion
            if node.right is None and node.left is None:
                self.root = None
            elif node.right is None:
                self.root = node.left
            elif node.left is None:
                self.root = node.right
            else:
                leftmost, parent = node.right, node
                while leftmost.left is not None:
                    leftmost, parent = leftmost.left, leftmost
                if leftmost == node.right:
                    leftmost.left = node.left
                else:
                    if leftmost.right is not None:
                        parent.left = leftmost.right
                    else:
                        parent.left = None   
                    leftmost.left, leftmost.right = node.left, node.right
                self.root = leftmost
        else:
            # other deletions
            if node.right is None and node.left is None:
                if last.right == node:
                    last.right = None
                if last.left == node:
                    last.left = None
            elif node.right is None:
                if last.right == node:
                    last.right = node.left
                if last.left == node:
                    last.left = node.left
            elif node.left is None:
                if last.right == node:
                    last.right = node.right
                if last.left == node:
                    last.left = node.right
            else:
                leftmost, parent = node.right, node
                while leftmost.left is not None:
                    leftmost, parent = leftmost.left, leftmost
                if node.right is not leftmost:
                    if leftmost.right is not None:
                        parent.left = leftmost.right
                    else:
                        parent.left = None
                    leftmost.left, leftmost.right = node.left, node.right
                    if last.right == node:
                        last.right = leftmost
                    elif last.left == node:
                        last.left = leftmost
                else:
                    if last.right == node:
                        last.right = leftmost
                    elif last.left == node:
                        last.left = leftmost
                    leftmost.left = node.left
        self.fix_prio(self.root)
```

<!-- grade_id: cell-20, points: 6, grade: true, solution: false -->
```python
b = DynamicTreap()
b.insert(6) 
b.insert(2)
b.insert(5)
b.insert(3)
b.insert(8)
b.insert(4)
b.insert(1)
b.insert(7)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
b.insert(5)
b.insert(4)
b.insert(7)
b.insert(5)
b.insert(7)
b.insert(7)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
```

Output:
```
--- preOrderTraverse ---
key 6 priority 0
key 2 priority 0
key 1 priority 0
key 5 priority 0
key 3 priority 0
key 4 priority 0
key 8 priority 0
key 7 priority 0
--- inOrderTraverse ---
key 1 priority 0
key 2 priority 0
key 3 priority 0
key 4 priority 0
key 5 priority 0
key 6 priority 0
key 7 priority 0
key 8 priority 0
--- postOrderTraverse ---
key 1 priority 0
key 4 priority 0
key 3 priority 0
key 5 priority 0
key 2 priority 0
key 7 priority 0
key 8 priority 0
key 6 priority 0
--- preOrderTraverse ---
key 7 priority 3
key 5 priority 2
key 4 priority 1
key 2 priority 0
key 1 priority 0
key 3 priority 0
key 6 priority 0
key 8 priority 0
--- inOrderTraverse ---
key 1 priority 0
key 2 priority 0
key 3 priority 0
key 4 priority 1
key 5 priority 2
key 6 priority 0
key 7 priority 3
key 8 priority 0
--- postOrderTraverse ---
key 1 priority 0
key 3 priority 0
key 2 priority 0
key 4 priority 1
key 6 priority 0
key 5 priority 2
key 8 priority 0
key 7 priority 3
```

<!-- grade_id: cell-21, points: 6, grade: true, solution: false -->
```python
b.remove(5)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
b.insert(5)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
```

Output:
```
--- preOrderTraverse ---
key 7 priority 3
key 4 priority 1
key 2 priority 0
key 1 priority 0
key 3 priority 0
key 6 priority 0
key 8 priority 0
--- inOrderTraverse ---
key 1 priority 0
key 2 priority 0
key 3 priority 0
key 4 priority 1
key 6 priority 0
key 7 priority 3
key 8 priority 0
--- postOrderTraverse ---
key 1 priority 0
key 3 priority 0
key 2 priority 0
key 6 priority 0
key 4 priority 1
key 8 priority 0
key 7 priority 3
--- preOrderTraverse ---
key 7 priority 3
key 4 priority 1
key 2 priority 0
key 1 priority 0
key 3 priority 0
key 6 priority 0
key 5 priority 0
key 8 priority 0
--- inOrderTraverse ---
key 1 priority 0
key 2 priority 0
key 3 priority 0
key 4 priority 1
key 5 priority 0
key 6 priority 0
key 7 priority 3
key 8 priority 0
--- postOrderTraverse ---
key 1 priority 0
key 3 priority 0
key 2 priority 0
key 5 priority 0
key 6 priority 0
key 4 priority 1
key 8 priority 0
key 7 priority 3
```

```python

```